---
title: "Tutorial for R for Data Science Ch 3 & 5"
output: learnr::tutorial
runtime: shiny_prerendered
---

```{r setup, include=FALSE}
library(learnr)
library(gradethis)
tutorial_options(exercise.checker = gradethis::grade_learnr)
knitr::opts_chunk$set(echo = FALSE)
library(tidyverse)
library(nycflights13)
```


## Section 3.3.1 

### 3.3.1 Easy DN

#### Exercise 1 DN

Write the R code for putting mpg dataset into a mpg dataframe

```{r D1, exercise=TRUE}

```

```{r D1-solution, include=FALSE}
mpg<- data.frame(mpg)
```

```{r D1-check}
gradethis::grade_this_code()
```

#### Exercise 2 DN

Write the R code for Viewing the mpg dataframe
```{r D2, exercise=TRUE}

```

```{r D2-solution}
View(mpg)
```

```{r D2-check}
gradethis::grade_this_code()
```

### 3.2.4 EASY WC
####  BLANK

## section 3.3.1 KA

### 3.3.1 Easy KA

write the code that will load tidyverse:

```{r two-plus-two, exercise=TRUE}

```
```{r two-plus-two-hint-1}
"first part is library()"

```
```{r two-plus-two-hint-2}
"you need to insert the word tidyverse"
```

```{r two-plus-two-solution}
library(tidyverse)
```
```{r two-plus-two-check}
gradethis::grade_code()
```

### 3.3.1 Easy KA

if you want learn more about the variables in mpg which code should you run

```{r two-plus-three, exercise=TRUE}

```
```{r two-plus-three-hint-1}
"use the ? sign"
```

```{r two-plus-three-solution}
?mpg
```

```{r two-plus-three-check}
gradethis:: grade_code()
```





###  3.3.1 Intermediate DN

#### Exercise 12 DN
Refer to mpg data set.  

Plot a geom_point graph with displ on x-axis, hwy on y-axis, and map the shape of points to drv
```{r D12, exercise=TRUE}
ggplot(    ) + 
  geom_point(    )
```

```{r D12-hint}
"Mapping the shape of points to drv means that the argument should be included inside aesthetic"
```

```{r D12-solution}
ggplot(data = mpg) + 
  geom_point(mapping = aes(x = displ, y = hwy, shape= drv))
```

```{r D12-check}
gradethis::grade_this_code()
```


#### Exercise 13 DN

```{r D13}
quiz(
  question("Which of the following are considered asethetics in geom_point ?",
    answer("fill", correct = TRUE),
    answer("shape", correct = TRUE),
    answer("size", correct= TRUE),
    answer("length")
  )
)
```

### 3.3.1 Intermediate KA

<!-- ![picture](images/Rplot22.png) -->
<img src="images/Rplot22.png" alt="drawing" style="width:400px;"/>

enter the code that creates the above graph

```{r two-plus-four, exercise=TRUE}

```

```{r two-plus-four-hint-1}
"you need to use both ggplot() and ggplot_point()"
```

```{r two-plus-four-hint-2}
"ggplot(data = mpg) + 
  geom_point(mapping = aes(x= , y= ))"
```

```{r two-plus-four-solution}
ggplot(data = mpg) + 
  geom_point(mapping = aes(x = displ, y = hwy))
```

```{r two-plus-four-check}
gradethis:: grade_code()
```

### 3.3.1 Intermediate KA

<!-- ![picture](images/Rplot23.png) -->
<img src="images/Rplot23.png" alt="drawing" style="width:400px;"/>

complete the code to create the above graph

```{r two-plus-five, exercise=TRUE}

ggplot(data = mpg) + 
  geom_point(mapping = aes(x = displ, y = hwy, ))
             
```

```{r two-plus-five-hint-1}
"missing part is size"
```

```{r two-plus-five-solution}
ggplot(data = mpg) + 
  geom_point(mapping = aes(x = displ, y = hwy, size = class))
```
```{r two-plus-five-check}
gradethis:: grade_code()
```

### 3.3.1 Advanced DN  

#### Exercise 23 DN

Divide mpg dataframe into 2 half and plot 2 separate geom_point graphs respectively with displ on x-axis and hwy on y-axis  
**HINT**:"spliting the mpg dataframe requires floor(), set.seed(), sample(), seq_len()"

```{r D23, exercise=TRUE}
half_size <- floor(0.5*  )
set.seed(123123)
mpg_half <- sample(seq_len(   ), size= half_size)

mpg_1 <- mpg[mpg_half, ]
mpg_2 <- mpg[-mpg_half, ]

ggplot(data= mpg_1) +
  geom_point(  )
ggplot(data=mpg_2) +
  geom_point(  )
```

```{r D23-solution}
half_size <- floor(0.5*nrow(mpg))
set.seed(123123)
mpg_half <- sample(seq_len(nrow(mpg)),size= half_size)

mpg_1 <- mpg[mpg_half, ]
mpg_2 <- mpg[-mpg_half, ]

ggplot(data= mpg_1) +
  geom_point(mapping = aes(x= displ, y = hwy))
ggplot(data=mpg_2) +
  geom_point(mapping = aes(x= displ, y = hwy))
```

```{r D23-check}
gradethis::grade_this_code()
```

### 3.3.1 Challenging KA

<!-- ![picture](images/Rplot24.png) -->
<img src="images/Rplot24.png" alt="drawing" style="width:400px;"/>


run the code that creates the graphs above

```{r two-plus-six, exercise=TRUE}

```
```{r two-plus-six-hint-1}
"you need to add shape = class"
```
```{r two-plus-six-solution}
ggplot(data = mpg) + 
  geom_point(mapping = aes(x = displ, y = hwy, shape = class))
```
```{r two-plus-six-check}
gradethis:: grade_code()
```

########################################
## Section 3.6.1 DN

### 3.6.1 Easy DN

#### Exercise 3 DN

```{r D3}
quiz(
  question("Which of the following are considered variables in facet_wrap ?",
    answer("rows"),
    answer("cols"),
    answer("nrow", correct= TRUE),
    answer("ncol", correct= TRUE)
  )
)
```

#### Exercise 4 DN

```{r D4}
quiz(
  question("Which of the following are considered variables in facet_grid ?",
    answer("rows", correct= TRUE),
    answer("cols", correct= TRUE),
    answer("nrow"),
    answer("ncol")
  )
)
```

### Easy 3.6.1 KA
```{r quiz1a}
question("which is geom?",
  answer("A geom is the geometrical object that a plot uses to represent data", correct = TRUE),
  answer("is a german car"),
  answer("is a programing language")
)
```

### 3.6.1 Easy KA
<!-- ![picture](images/Rplot28.png) -->
<img src="images/Rplot28.png" alt="drawing" style="width:400px;"/>


run the code to create the above graph using ggplot()


```{r two-plus-eight, exercise=TRUE}

```

```{r two-plus-eight-solution}
ggplot(data = mpg) + 
  geom_point(mapping = aes(x = displ, y = hwy))
```

```{r two-plus-eight-check}
gradethis:: grade_code()
```





### 3.6.1 Intermediate DN

#### Exercise 14 DN

Recreate the R code necessary to generate the following graphs
```{r D14-graph}
ggplot(data=mpg, mapping = aes(x= displ, y= hwy)) +
  geom_point() +
  geom_line()
```

```{r D14, exercise=TRUE}
ggplot(data=mpg, mapping = aes(   ,   )) +
  geom_point() +
  geom_line()
```

```{r D14-hint}
"You would need 2 geom graphs"
```

```{r D14-solution}
ggplot(data=mpg, mapping = aes(x= displ, y= hwy)) +
  geom_point() +
  geom_line()
```

```{r D14-check}
gradethis::grade_this_code()
```

#### Exercise 15 DN

Recreate the R code necessary to generate the following graphs
```{r D15-graph}
ggplot() +
  geom_point(data=mpg, mapping= aes(x=displ, y= hwy, color= drv)) +
  geom_line(data=mpg, mapping = aes(x=displ, y= hwy, color= drv))

```

```{r D15, exercise=TRUE}
ggplot() +
  geom_point() +
  geom_line()

```

```{r D15-hint}
"You should map the color of points to drv variable"
```

```{r D15-solution}
ggplot() +
  geom_point(data=mpg, mapping= aes(x=displ, y= hwy, color= drv)) +
  geom_line(data=mpg, mapping = aes(x=displ, y= hwy, color= drv))
```

```{r D15-check}
gradethis::grade_this_code()
```  

### 3.6.1 Intermediate KA

<!-- ![picture](images/Rplot29.png) -->
<img src="images/Rplot29.png" alt="drawing" style="width:400px;"/>

fix the code to create the graph above

```{r two-plus-nine, exercise=TRUE}
ggplot(data = mpg) + 
  geom_point(mapping = aes(x = displ, y = hwy, linetype = drv))
```
```{r two-plus-nine-hint-1}
"try geom _smooth() rather than geom_point()"
```

```{r two-plus-nine-solution}
ggplot(data = mpg) + 
  geom_smooth(mapping = aes(x = displ, y = hwy, linetype = drv))
```

```{r two-plus-nine-check}
gradethis:: grade_code()

```

### 3.6.1 Intermediate KA

<!-- ![picture](images/Rplot30.png) -->
<img src="images/Rplot30.png" alt="drawing" style="width:400px;"/>

create the above graph

```{r two-plus-ten, exercise=TRUE}

```

```{r two-plus-ten-hint-1}
"ggplot(data = mpg, mapping = aes(x = displ, y = hwy)) + "
```

```{r two-plus-ten-solution}
ggplot(data = mpg, mapping = aes(x = displ, y = hwy)) + 
  geom_point() + 
  geom_smooth()
```

```{r two-plus-ten-check}
gradethis::grade_code()

```



### Advanced 3.6.1 DN

#### Exercise 24 DN

Refer to the mpg dataset.  

Write the R code to plot 1 geom_point graph and 1 geom_smooth graph, with color=class, excluding all 2seater cars  

```{r D24, exercise=TRUE}
ggplot(data = mpg, mapping = aes(x = displ, y = hwy)) + 
   geom_point(data =   ) + 
   geom_smooth(data = , se=FALSE)
```

```{r D24-hint}
"Remind yourself how you utilized filter() function earlier "
```

```{r D24-solution}
ggplot(data = mpg, mapping = aes(x = displ, y = hwy)) + 
   geom_point(data = filter(mpg, class != "2seater"), mapping = aes(color = class)) + 
   geom_smooth(data = filter(mpg, class != "2seater"), se= FALSE)
```

```{r D24-check}
gradethis::grade_this_code()
```

### 3.6.1 Challenging KA

<!-- ![picture](images/Rplot31.png) -->
<img src="images/Rplot31.png" alt="drawing" style="width:400px;"/>

finish the code to create the above graph

```{r two-plus-zero, exercise=TRUE}
ggplot(data = mpg, mapping = aes(x = displ, y = hwy)) + 
  geom_point(mapping = aes(color = class)) + 
  geom_smooth(
```

```{r two-plus-zero-hint-1}
"you have to use the filter()"

```

```{r two-plus-zero-solution}
ggplot(data = mpg, mapping = aes(x = displ, y = hwy)) + geom_point(mapping = aes(color = class)) + geom_smooth(data = filter(mpg, class =="subcompact"), se = FALSE)  
```

```{r two-plus-zero-check}
gradethis:: grade_code()

```



## Section 5.2.4 DN

### Easy 5.2.4 DN

#### Exercise 5 DN

Refer to flights dataset.
Find all flights that arrived on Christmas Eve (24th December)

```{r D5, exercise=TRUE}
filter(flights,     &     )
```

```{r D5-hint}
"Experiment with the filter() function"
```

```{r D5-solution}
filter(flights, month == 12 & day == 24)
```

```{r D5-check}
gradethis::grade_this_code()
```

## Section 5.2.4 KA

### 5.2.4 Easy KA

Write the code to find the flights that had an arrival delay of 60 minutes or more.

```{r two-plus-8a, exercise=TRUE}

```
```{r two-plus-8a-hint-1}
"filter(flights, arr_delay >="
```

```{r two-plus-8a-solution}
filter(flights, arr_delay >= 60)
```
```{r two-plus-8a-check}
gradethis::grade_code()
```



### Intermediate 5.2.4 DN  

#### Exercise 16 DN  

Refer to flights dataset.
Find all flights with missing dep_delay, and arr_delay  

```{r D16, exercise=TRUE}
filter(flights,     &     )
```


```{r D16-hint}
"Experiment with is.na() and filter() function"
```

```{r D16-solution}
ggplot() +
  geom_point(data=mpg, mapping= aes(x=displ, y= hwy, color= drv)) +
  geom_line(data=mpg, mapping = aes(x=displ, y= hwy, color= drv))
```

```{r D16-check}
gradethis::grade_this_code()
```

### 5.2.4 Intermediate KA

wirte the code to find flights that departed between January and August

```{r two-plus-threea, exercise=TRUE}

```
```{r two-plus-threea-hint-1}
filter(flights, month >= , month <= )
```
```{r two-plus-threea-solution}
filter(flights, month >= 7, month <= 9)
```
```{r two-plus-threea-check}
gradethis::grade_code()
```



### Advanced 5.2.4 DN  

#### Exercise 25 DN  

Refer to flights dataset.
Find all flights that belong to Mesa Airlines, Hawaiian Airlines, and Alaska Airlines

```{r D25-hint}
"type airplanes in the console"
```

```{r D25, exercise=TRUE}
filter(flights, carrier ==    )
```

```{r D25-solution}
filter(flights, carrier == c("YV","HA","AS"))
```

```{r D25-check}
gradethis::grade_this_code()
```  

### 5.2.4 Challenging KA

find the flights that Departed between midnight and 6am (inclusive).

```{r two-plus-fourb, exercise=TRUE}

```
```{r two-plus-fourb-solution}
summary(flights$dep_time)
filter(flights, dep_time <= 600 | dep_time == 2400)
c(600, 1200, 2400) %% 2400
filter(flights, dep_time %% 2400 <= 600)
```

```{r two-plus-fourb-check}
gradethis::grade_code()
```



#############################################
## Section 5.3.1 DN

### Easy 5.3.1 DN

#### Exercise 6 DN

```{r D6}
quiz(
  question("Which of the following are considered arguments in arrange() ?",
    answer("data"),
    answer(".data", correct= TRUE),
    answer("by_group"),
    answer(".by_group", correct= TRUE)
  )
)
```

#### Exercise 7 DN  

```{r D7}
quiz(
  question("Using arrange() function, where are NA values located ?",
    answer("top of the data"),
    answer("randomized"),
    answer("end of the data", correct= TRUE)
  )
)
```

## section 5.3.1 KA

### 5.3.1 Easy KA

Sort flights to find the most delayed flights. 

```{r two-plus-1a, exercise=TRUE}

```
```{r two-plus-1a-hint-1}
"use desc( ))"
```
```{r two-plus-1a-solution}
arrange(flights, desc(dep_delay))
```
```{r two-plus-1a-check}
gradethis::grade_code()
```



### Intermediate 5.3.1 DN  

#### Exercise 17 DN 

Refer to flights dataset.
Arrange flights based on descending dep_delay  

```{r D17, exercise=TRUE}
arrange(  ,   )
```

```{r D17-hint}
"Experiment with arrange() function"
```

```{r D17-solution}
arrange(flights, desc(dep_delay))
```

```{r D17-check}
gradethis::grade_this_code()
```

#### Exercise 18 DN 

Refer to flights dataset.
Arrange flights based on ascending distance

```{r D18, exercise=TRUE}
arrange(    ,    )
```

```{r D18-hint}
"Ascending order is the default for arrange()"
```

```{r D18-solution}
arrange(flights, distance)
```

```{r D18-check}
gradethis::grade_this_code()
```


### 5.3.1 Intermediate KA

write the code to find the Flights that left the earliest.

```{r two-plus-2a, exercise=TRUE}

```
```{r two-plus-2a-hint-1}
"the code looks similiar to the answer of question 1"
```
```{r two-plus-2a-solution}
arrange(flights, dep_delay)
```
```{r two-plus-2a-check}
gradethis::grade_code()
```



### Advanced 5.3.1 DN

#### Exercise 27 DN

Make a new column in flights dataset called speed where speed = distance/air_time. Arrange flights based on descending speed
```{r D27, exercise=TRUE}
flights <- flights %>%
  mutate(speed =   /   )
  arrange(  ,  )
```

```{r D27-hint}
"Experiment with contains() function in dplyr package "
```

```{r D27-solution}
flights <- flights %>%
  mutate(speed = distance/air_time)
  arrange(flights, desc(speed))
```

```{r D27-check}
gradethis::grade_this_code()
```

### 5.3.1 Challenging KA

show top 10 flights that traveled the farthest and top 10 that traveled the shortest?

```{r two-plus-3a, exercise=TRUE}

```
```{r two-plus-3a-hint-1}
"The answer follows the same idea as the answer for number 1 and 2"
```
```{r two-plus-3a-solution}
arrange(flights, desc(distance))
arrange(flights, distance)
```
```{r two-plus-3a-check}
gradethis::grade_code()
```



#######################################


## Section 5.4.1 DN

### Easy 5.4.1 DN  

#### Exercise 8 DN  

Refer to flights dataset.
Rename air_time column to airtime 

```{r D8, exercise=TRUE}
rename(flights,   =     )
```

```{r D8-hint}
"Experiment with the rename() function"
```


```{r D8-solution}
rename(flights, airtime = air_time)
```

```{r D8-check}
gradethis::grade_this_code()
```


## Section 5.4.1 KA

### 5.4.1 Easy KA

Use any_of () function to find dep_delay and dep_time vectors. 

```{r two-plus-1b, exercise=TRUE}

```

```{r two-plus-1b-hint-1}
"select(flights, any_of(c()))"
```

```{r two-plus-1b-solution}
select(flights, any_of(c("dep_time", "dep_delay", "arr_time", "arr_delay")))
```

```{r two-plus-1b-check}
gradethis::grade_code()
```




### Intermediate 5.4.1 DN  

#### Exercise 19 DN 

Refer to flights dataset.
Select all columns except year, month, day column  

```{r D19, exercise=TRUE}
select(flights,    )
```

```{r D19-hint}
"You can either list all other columns beside year, month, day or use `!` "
```

```{r D19-solution}
select(flights, !c(year, month, day))
```

```{r D19-check}
gradethis::grade_this_code()
```  

### 5.4.1 Intermediat KA

find the vectors dep_time, dep_delay, arr_time, arr_delay without using dep_time, dep_delay, arr_time, arr_delay in the function,, use the starts_with()function.

```{r two-plus-2b, exercise=TRUE}

```
```{r two-plus-2b-solution}
select(flights, starts_with("dep_"), starts_with("arr_"))
```
```{r two-plus-2b-check}
gradethis::grade_code()
```




### Advanced 5.4.1 DN  

#### Exercise 26 DN 
Refer to flights dataset.
Select all columns not containing the word "time"

```{r D26, exercise=TRUE}
select(flights,    )
```

```{r D26-hint}
"Experiment with contains() function in dplyr package "
```

```{r D26-solution}
select(flights, !contains("time"))
```

```{r D26-check}
gradethis::grade_this_code()
```


### 5.4.1 Challenging KA

use the !! operator in a function to find the vectors of "dep_time", "dep_delay", "arr_time", "arr_delay".

```{r two-plus-3b, exercise=TRUE}

```
```{r two-plus-3b-hint-1}
 "Specify the names of the variables with a character vector and use the bang-bang operator (use c as a name for the variables) (!!)."
```
```{r two-plus-3b-solution}
variables <- c("dep_time", "dep_delay", "arr_time", "arr_delay")
select(flights, !!variables)
```
```{r two-plus-3b-check}
gradethis::grade_code()
```

########################################################


## Section 5.5.2 DN

### Easy 5.5.2 DN

#### Exercise 9 DN

```{r D9}
quiz(
  question("Which of the following arguments control where new columns should appear in mutate() ?",
    answer(".data"),
    answer(".keep"),
    answer(".before"),
    answer(".after", correct= TRUE)
  )
)
```

#### Exercise 10 DN 

```{r D10}
quiz(
  question("Using mutate(), What will happen to NULL values in output column ?",
    answer("removed", correct= TRUE),
    answer("retained"),
    answer("changed to NA"),
    answer("changed to logical(0)")
  )
)
```

## Section 5.5.2 KA

### 5.5.2 Easy KA

```{r quiz2b}
question("what would the following code 1504 %/% 100 give",
  answer("[1] 15", correct = TRUE),
  answer("[1] 4"),
  answer("150,400"),
  answer("15.04")
)
```




### Intermediate 5.5.2 DN

#### Exercise 20 DN  

Refer to flights dataset.
Add a new column called air_time_hours by converting air_time from minutes to hours

```{r D20, exercise=TRUE}
mutate(    ,    )
```

```{r D20-hint}
"Experiment with mutate() function "
```

```{r D20-solution}
mutate(flights, air_time_hours = air_time/60)
```

```{r D20-check}
gradethis::grade_this_code()
```

#### Exercise 21 DN  

Refer to flights dataset.
Add a new column called distance_km by converting distance from miles to km (1 mile = 1.61 km)  

```{r D21, exercise=TRUE}
mutate(    ,    )
```

```{r D21-hint}
"Experiment with mutate() function "
```

```{r D21-solution}
mutate(flights, distance_km = distance * 1.61)
```

```{r D21-check}
gradethis::grade_this_code()
```


### 5.5.2  Intermediate KA

```{r quiz1c}
question(" what does the following give us 1:3 + 1:10",
  answer("[1]  2  4  6  5  7  9  8 10 12 11", correct = TRUE),
  answer("[1]  2  4  6  5  7  9  9 10 12 11"),
  answer("[1]  2  3  4  5  6  7  8 9 10 11"),
  answer("gives us nothing")
)
```




### Advanced 5.5.2 DN  

#### Exercise 28 DN   

Refer to flights dataset.
Rank the 10 fastest flights and 10 slowest flights  

```{r D28-hint}
"Experiment with top_n() function"
```

```{r D28, exercise=TRUE}
flights <- mutate(flights, speed= distance/air_time)
arrange(  ,  ) %>% 
arrange(  ,  ) %>% 
```

```{r D28-solution}
flights <- mutate(flights, speed=distance/air_time)
arrange(flights, desc(speed)) %>% top_n(10)
arrange(flights, speed) %>% top_n(10)
```

```{r D28-check}
gradethis::grade_this_code()
```
### 5.5.2 Challenging KA

Compare dep_time, sched_dep_time, and dep_delay. 

```{r two-plus-1d, exercise=TRUE}

```

```{r two-plus-1d-solution}
flights_deptime <-
  mutate(flights,
    dep_time_min = (dep_time %/% 100 * 60 + dep_time %% 100) %% 1440,
    sched_dep_time_min = (sched_dep_time %/% 100 * 60 +
      sched_dep_time %% 100) %% 1440,
    dep_delay_diff = dep_delay - dep_time_min + sched_dep_time_min
  )
filter(flights_deptime, dep_delay_diff != 0)
```

```{r two-plus-1d-check}
gradethis::grade_code()
```

###############################################
## Section 5.6.7 DN

### Easy 5.6.7 DN  

#### Exercise 11 DN  
Refer to flights dataset.
Find the mean air time for each month

```{r D11, exercise = TRUE}
flights %>% 
  group_by( ) %>%
  summarise(mean_air_time = mean( ))
```

```{r D11-hint}
"NAs should be removed in calculating the mean air time"
```

```{r D11-solution}
flights %>% 
  group_by(year, month) %>%
  summarise(mean_air_time = mean(air_time, na.rm = TRUE))
```

```{r D11-check}
gradethis::grade_this_code()
```

## Section 5.6.7 KA

### 5.6.7 Easy KA

```{r quiz1e}
question("which delay may prove costlier to customer?",
  answer("arrival delay", correct = TRUE),
  answer("departure delay")
)
```




### Intermediate 5.6.7 DN  

#### Exercise 22 DN  

Refer to flights dataset.
Find the median air time for each day of the year

```{r D22, exercise=TRUE}
flights %>% 
  group_by(   ) %>%
  summarise(median_air_time = median(   ))
```

```{r D22-hint}
"NAs should be removed in calculating the mean air time "
```

```{r D22-solution}
flights %>% 
  group_by(year, month, day) %>%
  summarise(median_air_time = median(air_time, na.rm = TRUE))
```

```{r D22-check}
gradethis::grade_this_code()
```

### 5.6.7 Intermediate KA


write the code to show Which carrier has the worst delays(arr_delay)?

```{r two-plus-foura, exercise=TRUE}

```
```{r two-plus-foura-hint-1}
"flights %>% group_by(carrier, dest) %>% summarise(n())"
```
```{r two-plus-foura-solution}
flights %>%
  group_by(carrier) %>%
  summarise(arr_delay = mean(arr_delay, na.rm = TRUE)) %>%
  arrange(desc(arr_delay))
filter(airlines, carrier == "F9")
```
```{r two-plus-foura-check}
gradethis::grade_code()
```





### Advanced 5.6.7 DN  

#### Exercise 29 (5.6.7)  

Refer to flights dataset.
Find the first departure and last arrival for each day of the year, and plot the following graph  

```{r D29-graph}
flights %>%
  filter(!is.na(dep_time)&!is.na(arr_time)) %>%
  group_by(year, month, day) %>%
  summarise(first_dep = first(dep_time),
            last_arr = last(arr_time)) %>%
  ggplot(aes(first_dep,last_arr)) +
  geom_smooth()
```

```{r D29, exercise=TRUE}
flights %>%
  filter(    &    ) %>%
  group_by(    ) %>%
  summarise(first_dep = first(   ),
            last_arr  =  last(   )) %>%
  ggplot(aes(first_dep,last_arr)) +
  geom_smooth()
```

```{r D29-hint}
"Refer back to our prevous exercises which use is.na(), group_by(), and summarise()"
```

```{r D29-solution}
flights %>%
  filter(!is.na(dep_time)&!is.na(arr_time)) %>%
  group_by(year, month, day) %>%
  summarise(first_dep = first(dep_time),
            last_arr = last(arr_time)) %>%
  ggplot(aes(first_dep,last_arr)) +
  geom_smooth()
```

```{r D29-check}
gradethis::grade_this_code()
```

### 5.6.7 Challenging KA

What time of day should you fly if you want to avoid delays as much as possible?
 
 ```{r two-plus-fivec, exercise=TRUE}

```
```{r two-plus-fivec-solution}
flights %>%
  group_by(hour) %>%
  summarise(arr_delay = mean(arr_delay, na.rm = TRUE)) %>%
  arrange(arr_delay)
```
```{r two-plus-fivec-check}
gradethis::grade_code()
```




####################################







